---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.
# Test the comment

`@instructions`
[www.facebook.com]()

`@hint`


`@pre_exercise_code`
```{python}
import pandas as pd
```

`@sample_code`
```{python}

```

`@solution`
```{python}
df = pd.DataFrame({'a':2, 
                   'b':4}
                 )
```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 011bbb8da8
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
